[1540955128] SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;1;DISK CRITICAL - /run/user/1000/gvfs is not accessible: Permission denied
[1540955188] SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;2;DISK CRITICAL - /run/user/1000/gvfs is not accessible: Permission denied
[1540955228] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;1;connect to address 127.0.0.1 and port 22: Connection refused
[1540955248] SERVICE ALERT: localhost;Disk Space;CRITICAL;SOFT;3;DISK CRITICAL - /run/user/1000/gvfs is not accessible: Permission denied
[1540955288] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;2;connect to address 127.0.0.1 and port 22: Connection refused
[1540955308] SERVICE ALERT: localhost;Disk Space;CRITICAL;HARD;4;DISK CRITICAL - /run/user/1000/gvfs is not accessible: Permission denied
[1540955348] SERVICE ALERT: localhost;SSH;CRITICAL;SOFT;3;connect to address 127.0.0.1 and port 22: Connection refused
[1540955406] SERVICE ALERT: localhost;SSH;CRITICAL;HARD;4;connect to address 127.0.0.1 and port 22: Connection refused
[1541028039] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;1;HTTP WARNING: HTTP/1.1 404 Not Found - 453 bytes in 0.001 second response time
[1541028099] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;2;HTTP WARNING: HTTP/1.1 404 Not Found - 453 bytes in 0.001 second response time
[1541028159] SERVICE ALERT: localhost;HTTP;WARNING;SOFT;3;HTTP WARNING: HTTP/1.1 404 Not Found - 453 bytes in 0.001 second response time
[1541028219] SERVICE ALERT: localhost;HTTP;WARNING;HARD;4;HTTP WARNING: HTTP/1.1 404 Not Found - 453 bytes in 0.000 second response time
[1541083103] SERVICE ALERT: localhost;HTTP;OK;HARD;4;HTTP OK: HTTP/1.1 302 Found - 546 bytes in 0.002 second response time
[1541083403] SERVICE ALERT: localhost;HTTP;CRITICAL;SOFT;1;HTTP CRITICAL: HTTP/1.1 500 500 Service unavailable (with message) - 520 bytes in 2.484 second response time
[1541083463] SERVICE ALERT: localhost;HTTP;CRITICAL;SOFT;2;HTTP CRITICAL: HTTP/1.1 500 500 Service unavailable (with message) - 520 bytes in 0.154 second response time
[1541083523] SERVICE ALERT: localhost;HTTP;CRITICAL;SOFT;3;HTTP CRITICAL: HTTP/1.1 500 500 Service unavailable (with message) - 520 bytes in 0.095 second response time
[1541083583] SERVICE ALERT: localhost;HTTP;CRITICAL;HARD;4;HTTP CRITICAL: HTTP/1.1 500 500 Service unavailable (with message) - 520 bytes in 0.061 second response time
[1541083883] SERVICE ALERT: localhost;HTTP;OK;HARD;4;HTTP OK: HTTP/1.1 200 OK - 8998 bytes in 0.421 second response time